# Problem 2: DNA Sequence Analysis

This problem is about strings and regular expressions. It has four (4) exercises, numbered 0-3. They are worth a total of ten (10) points.

```python
import re # You'll need this module
```

## DNA Sequence Analysis

Your friend is a biologist who is studying a particular DNA sequence. The sequence is a string built from an alphabet of four possible letters, `A`, `G`, `C`, and `T`. Biologists refer to each of these letters a _base_.

Here is an example of a DNA fragment as a string of bases.

```python
dna_seq = 'ATGGCAATAACCCCCCGTTTCTACTTCTAGAGGAGAAAAGTATTGACATGAGCGCTCCCGGCACAAGGGCCAAAGAAGTCTCCAATTTCTTATTTCCGAATGACATGCGTCTCCTTGCGGGTAAATCACCGACCGCAATTCATAGAAGCCTGGGGGAACAGATAGGTCTAATTAGCTTAAGAGAGTAAATCCTGGGATCATTCAGTAGTAACCATAAACTTACGCTGGGGCTTCTTCGGCGGATTTTTACAGTTACCAACCAGGAGATTTGAAGTAAATCAGTTGAGGATTTAGCCGCGCTATCCGGTAATCTCCAAATTAAAACATACCGTTCCATGAAGGCTAGAATTACTTACCGGCCTTTTCCATGCCTGCGCTATACCCCCCCACTCTCCCGCTTATCCGTCCGAGCGGAGGCAGTGCGATCCTCCGTTAAGATATTCTTACGTGTGACGTAGCTATGTATTTTGCAGAGCTGGCGAACGCGTTGAACACTTCACAGATGGTAGGGATTCGGGTAAAGGGCGTATAATTGGGGACTAACATAGGCGTAGACTACGATGGCGCCAACTCAATCGCAGCTCGAGCGCCCTGAATAACGTACTCATCTCAACTCATTCTCGGCAATCTACCGAGCGACTCGATTATCAACGGCTGTCTAGCAGTTCTAATCTTTTGCCAGCATCGTAATAGCCTCCAAGAGATTGATGATAGCTATCGGCACAGAACTGAGACGGCGCCGATGGATAGCGGACTTTCGGTCAACCACAATTCCCCACGGGACAGGTCCTGCGGTGCGCATCACTCTGAATGTACAAGCAACCCAAGTGGGCCGAGCCTGGACTCAGCTGGTTCCTGCGTGAGCTCGAGACTCGGGATGACAGCTCTTTAAACATAGAGCGGGGGCGTCGAACGGTCGAGAAAGTCATAGTACCTCGGGTACCAACTTACTCAGGTTATTGCTTGAAGCTGTACTATTTTAGGGGGGGAGCGCTGAAGGTCTCTTCTTCTCATGACTGAACTCGCGAGGGTCGTGAAGTCGGTTCCTTCAATGGTTAAAAAACAAAGGCTTACTGTGCGCAGAGGAACGCCCATCTAGCGGCTGGCGTCTTGAATGCTCGGTCCCCTTTGTCATTCCGGATTAATCCATTTCCCTCATTCACGAGCTTGCGAAGTCTACATTGGTATATGAATGCGACCTAGAAGAGGGCGCTTAAAATTGGCAGTGGTTGATGCTCTAAACTCCATTTGGTTTACTCGTGCATCACCGCGATAGGCTGACAAAGGTTTAACATTGAATAGCAAGGCACTTCCGGTCTCAATGAACGGCCGGGAAAGGTACGCGCGCGGTATGGGAGGATCAAGGGGCCAATAGAGAGGCTCCTCTCTCACTCGCTAGGAGGCAAATGTAAAACAATGGTTACTGCATCGATACATAAAACATGTCCATCGGTTGCCCAAAGTGTTAAGTGTCTATCACCCCTAGGGCCGTTTCCCGCATATAAACGCCAGGTTGTATCCGCATTTGATGCTACCGTGGATGAGTCTGCGTCGAGCGCGCCGCACGAATGTTGCAATGTATTGCATGAGTAGGGTTGACTAAGAGCCGTTAGATGCGTCGCTGTACTAATAGTTGTCGACAGACCGTCGAGATTAGAAAATGGTACCAGCATTTTCGGAGGTTCTCTAACTAGTATGGATTGCGGTGTCTTCACTGTGCTGCGGCTACCCATCGCCTGAAATCCAGCTGGTGTCAAGCCATCCCCTCTCCGGGACGCCGCATGTAGTGAAACATATACGTTGCACGGGTTCACCGCGGTCCGTTCTGAGTCGACCAAGGACACAATCGAGCTCCGATCCGTACCCTCGACAAACTTGTACCCGACCCCCGGAGCTTGCCAGCTCCTCGGGTATCATGGAGCCTGTGGTTCATCGCGTCCGATATCAAACTTCGTCATGATAAAGTCCCCCCCTCGGGAGTACCAGAGAAGATGACTACTGAGTTGTGCGAT'
print("=== Sequence (Number of bases: {}) ===\n\n{}".format(len(dna_seq), dna_seq))
```

Output:
```
=== Sequence (Number of bases: 2012) ===

ATGGCAATAACCCCCCGTTTCTACTTCTAGAGGAGAAAAGTATTGACATGAGCGCTCCCGGCACAAGGGCCAAAGAAGTCTCCAATTTCTTATTTCCGAATGACATGCGTCTCCTTGCGGGTAAATCACCGACCGCAATTCATAGAAGCCTGGGGGAACAGATAGGTCTAATTAGCTTAAGAGAGTAAATCCTGGGATCATTCAGTAGTAACCATAAACTTACGCTGGGGCTTCTTCGGCGGATTTTTACAGTTACCAACCAGGAGATTTGAAGTAAATCAGTTGAGGATTTAGCCGCGCTATCCGGTAATCTCCAAATTAAAACATACCGTTCCATGAAGGCTAGAATTACTTACCGGCCTTTTCCATGCCTGCGCTATACCCCCCCACTCTCCCGCTTATCCGTCCGAGCGGAGGCAGTGCGATCCTCCGTTAAGATATTCTTACGTGTGACGTAGCTATGTATTTTGCAGAGCTGGCGAACGCGTTGAACACTTCACAGATGGTAGGGATTCGGGTAAAGGGCGTATAATTGGGGACTAACATAGGCGTAGACTACGATGGCGCCAACTCAATCGCAGCTCGAGCGCCCTGAATAACGTACTCATCTCAACTCATTCTCGGCAATCTACCGAGCGACTCGATTATCAACGGCTGTCTAGCAGTTCTAATCTTTTGCCAGCATCGTAATAGCCTCCAAGAGATTGATGATAGCTATCGGCACAGAACTGAGACGGCGCCGATGGATAGCGGACTTTCGGTCAACCACAATTCCCCACGGGACAGGTCCTGCGGTGCGCATCACTCTGAATGTACAAGCAACCCAAGTGGGCCGAGCCTGGACTCAGCTGGTTCCTGCGTGAGCTCGAGACTCGGGATGACAGCTCTTTAAACATAGAGCGGGGGCGTCGAACGGTCGAGAAAGTCATAGTACCTCGGGTACCAACTTACTCAGGTT
```

In this problem, you will help your friend analyze this sequence.

**Exercise 0** (2 point). Complete the function, `count_bases(s)`. It takes as input a DNA sequence as a string, `s`. It should compute the number of occurrences of each base (i.e., `'A'`, `'C'`, `'G'`, and `'T'`) in `s`. It should then return these counts in a dictionary whose keys are the bases.

<!-- grade_id: exercise_0, points: 0, grade: false, solution: true -->
```python
def count_bases(s):
    assert type(s) is str
    assert all([b in ['A', 'C', 'G', 'T'] for b in s])
    
    count_A = 0
    count_C = 0
    count_G = 0
    count_T = 0
    
    count_dict = {}
    
    for b in s:
        if b == 'A':
            count_A+=1
        elif b == 'C':
            count_C+=1
        elif b == 'G':
            count_G+=1
        elif b == 'T':
            count_T+=1

    count_dict['A']=count_A
    count_dict['C']=count_C
    count_dict['G']=count_G
    count_dict['T']=count_T
    
    return count_dict
```

The `Counter` class is part of the Python `collections` module, and it is used to count the occurrences of elements in an iterable, such as a list, tuple, or string. It's a convenient way to perform frequency counting and is especially useful when you need to analyze data or find the most common elements in a collection. 

 `Counter` creates a dictionary-like object where the keys are the unique elements in the iterable, and the values are their corresponding counts. You can then access the count of a specific element using square brackets, or you can use the `most_common` method to find the most common elements and their counts, which returns a list of tuples sorted by count in descending order.

```python
from collections import Counter

# Create a Counter object from an iterable (e.g., a list)
my_list = [1, 2, 2, 3, 3, 3, 4, 4, 4, 4]
counter = Counter(my_list)

print(counter)
print(type(counter))
```

Output:
```
Counter({4: 4, 3: 3, 2: 2, 1: 1})
<class 'collections.Counter'>
```

```python
dict(counter)
```

Output:
```
{1: 1, 2: 2, 3: 3, 4: 4}
```

```python
def count_bases(s):
    assert type(s) is str
    assert all([b in ['A', 'C', 'G', 'T'] for b in s])
    
    from collections import Counter
    return dict(Counter(s))
```

<!-- grade_id: exercise_0_test, points: 2, grade: true, solution: false -->
```python
# Test cell: `exercise_0_test`

base_counts = count_bases(dna_seq)
print("Your result:", base_counts)

assert type(base_counts) is dict, "`base_counts` is of type `{}`, not `dict`.".format(type(base_counts))
assert len(base_counts) <= 4, "There can be at most 4 bases."
for b, c in [('A', 501), ('C', 507), ('G', 508), ('T', 496)]:
    assert base_counts[b] == c, "Base '{}' has a count of {} when it should be {}.".format(b, base_counts[b], c)
    
print("\n(Passed!)")
```

Output:
```
Your result: {'A': 501, 'T': 496, 'G': 508, 'C': 507}

(Passed!)
```

**Enzyme "scissors."** Your friend is interested in what will happen to the sequence if she uses certain "restriction enzymes" to cut it. The enzymes work by scanning the DNA sequence from left to right for a particular pattern. It then cuts the DNA wherever it finds a match.

**A biologist's notation.** Your friend does not know about regular expressions. Instead, she uses a [special notation](https://en.wikipedia.org/wiki/Nucleic_acid_sequence) that other biologists use to describe base patterns. These are "extra letters" that have a special meaning.

For example, the special letter `N` denotes any base, i.e., any single occurrence of an `A`, `C`, `G`, or `T`. Therefore, when a biologist writes, `ANT`, that means `AAT`, `ACT`, `AGT`, or `ATT`.

Here is the complete set of special letters:

* `R`: Either `G` or `A`
* `Y`: Either `T` or `C`
* `K`: Either `G` or `T`
* `M`: Either `A` or `C`
* `S`: Either `G` or `C`
* `W`: Either `A` or `T`
* `B`: Anything but `A` (i.e., `G`, `T`, or `C`)
* `D`: Anything but `C`
* `H`: Anything but `G`
* `V`: Anything but `T`
* `N`: Anything, i.e., `A`, `C`, `G`, or `T`

---

**Exercise 1** (4 points). Given a string in the biologist's notation, complete the function `bio_to_regex(pattern_bio)` so that it returns an equivalent pattern in Python's regular expression language.

If your function is correct, then the following code would also work:

```python
  assert re.search(bio_to_regex('ANT'), 'AGATTA') is not None
```

That's because `ANT` matches `ATT`, which is contained in `AGATTA`.

n Python, the re.search() function is part of the re module (regular expressions module) and is used for searching a given pattern in a string. It returns a match object if the pattern is found in the string, and None if no match is found

match = re.search(pattern, text)

```python
pattern_bio = 'AGCCCTTTRAGGTTTAGC'

pattern_bio = pattern_bio.replace('R', '[GA]')

pattern_bio
```

Output:
```
'AGCCCTTT[GA]AGGTTTAGC'
```

In the string 'AGCCCTTT[GA]AGGTTTAGC', the portion inside square brackets '[GA]' represents a character class in Python regular expressions. Character classes are used to match any one character from a set of characters.

In this specific string:

'AGCCCTTT' matches the characters 'AGCCCTTT' exactly in that order.
'[GA]' is a character class that matches either 'G' or 'A', so it will match either 'G' or 'A' at this position in the string.
'AGGTTTAGC' matches the characters 'AGGTTTAGC' exactly in that order.
So, the string 'AGCCCTTT[GA]AGGTTTAGC' will match strings that have 'AGCCCTTT' followed by either 'G' or 'A' and then followed by 'AGGTTTAGC'. For example, it would match 'AGCCCTTTGAGGTTTAGC' and 'AGCCCTTTAAGGTTTAGC', but it wouldn't match 'AGCCCTTTTAGGTTTAGC' because there's no 'G' or 'A' between 'AGCCCTTT' and 'AGGTTTAGC' in that case.

```python
pattern_bio = 'AGCCCTTTHAGGTTTAGC'

pattern_bio = pattern_bio.replace('H', '[^G]')

pattern_bio
```

Output:
```
'AGCCCTTT[^G]AGGTTTAGC'
```

In the string 'AGCCCTTT[^A]AGGTTTAGC', the portion inside square brackets '[^A]' represents a negated character class in Python regular expressions. A negated character class matches any character that is not in the specified set of characters.

In this specific string:
- 'AGCCCTTT' matches the characters 'AGCCCTTT' exactly in that order.
- '[^A]' is a negated character class that matches any character that is not 'A'. So, it will match any character except 'A' at this position in the string.
- 'AGGTTTAGC' matches the characters 'AGGTTTAGC' exactly in that order.

So, the string 'AGCCCTTT[^A]AGGTTTAGC' will match strings that have 'AGCCCTTT' followed by any character except 'A' and then followed by 'AGGTTTAGC'. For example, it would match 'AGCCCTTTGAGGTTTAGC' and 'AGCCCTTTTAGGTTTAGC', but it wouldn't match 'AGCCCTTTAAGGTTTAGC' because there's an 'A' between 'AGCCCTTT' and 'AGGTTTAGC' in that case.

<!-- grade_id: exercise_1, points: 0, grade: false, solution: true -->
```python
def bio_to_regex(pattern_bio):
    pattern_bio = pattern_bio.replace('R', '[GA]')
    pattern_bio = pattern_bio.replace('Y', '[TC]')
    pattern_bio = pattern_bio.replace('K', '[GT]')
    pattern_bio = pattern_bio.replace('M', '[AC]')
    pattern_bio = pattern_bio.replace('S', '[GC]')
    pattern_bio = pattern_bio.replace('W', '[AT]')
    pattern_bio = pattern_bio.replace('B', '[^A]')
    pattern_bio = pattern_bio.replace('D', '[^C]')
    pattern_bio = pattern_bio.replace('H', '[^G]')
    pattern_bio = pattern_bio.replace('V', '[^T]')
    pattern_bio = pattern_bio.replace('N', '[ACGT]')
    return pattern_bio
```

```python
bio_map = {'R': '[GA]', 'Y': '[TC]', 'K': '[GT]', 'M': '[AC]', 'S': '[GC]',
               'W': '[AT]', 'B': '[^A]', 'D': '[^C]', 'H': '[^G]', 'V': '[^T]', 'N': '[ACGT]'}

pattern_bio = 'AGCCCTTTHAGGTTTAGC'

print(''.join(bio_map.get(letter, letter) for letter in pattern_bio))
```

Output:
```
AGCCCTTT[^G]AGGTTTAGC
```

```python
for letter in pattern_bio:
    print(bio_map.get(letter))#, letter))
```

Output:
```
None
None
None
None
None
None
None
None
[^G]
None
None
None
None
None
None
None
None
None
```

```python
string_ =''

for letter in pattern_bio:
    mapped_letter = bio_map.get(letter, letter)
    string_ += mapped_letter
    print(bio_map.get(letter, letter))
```

Output:
```
A
G
C
C
C
T
T
T
[^G]
A
G
G
T
T
T
A
G
C
```

```python
string_
```

Output:
```
'AGCCCTTT[^G]AGGTTTAGC'
```

```python
result_list = []

for letter in pattern_bio:
    mapped_letter = bio_map.get(letter, letter)
    result_list.append(mapped_letter)

result_string = ''.join(result_list)
result_string
```

Output:
```
'AGCCCTTT[^G]AGGTTTAGC'
```

```python
result_list
```

Output:
```
['A',
 'G',
 'C',
 'C',
 'C',
 'T',
 'T',
 'T',
 '[^G]',
 'A',
 'G',
 'G',
 'T',
 'T',
 'T',
 'A',
 'G',
 'C']
```

```python
def bio_to_regex(pattern_bio):
    
    bio_map = {'R': '[GA]', 'Y': '[TC]', 'K': '[GT]', 'M': '[AC]', 'S': '[GC]',
                'W': '[AT]', 'B': '[^A]', 'D': '[^C]', 'H': '[^G]', 'V': '[^T]', 'N': '[ACGT]'}


    return ''.join([bio_map.get(letter, letter) for letter in pattern_bio]) 

# For each letter, it looks up the corresponding regular expression character class from the bio_map 
# dictionary using the get method. If the letter is not found in the dictionary, it leaves it unchanged.
```

The function iterates over each letter in the input pattern_bio string using a list comprehension.

For each letter, it looks up the corresponding regular expression character class from the bio_map dictionary using the get method. If the letter is not found in the dictionary, it leaves it unchanged.

<!-- grade_id: exercise_1_test_0, points: 2, grade: true, solution: false -->
```python
# Test cell: `exercise_1_test_0`

assert re.search(bio_to_regex('ANT'), 'AGATTA') is not None
assert set(re.findall(bio_to_regex('ANTAAT'), dna_seq)) == {'ATTAAT', 'ACTAAT'}
assert set(re.findall(bio_to_regex('GCRWTG'), dna_seq)) == {'GCGTTG', 'GCAATG'}
assert len(re.findall(bio_to_regex('CDCHA'), dna_seq)) == 18

print("\n(Passed first group of tests!)")
```

Output:
```
(Passed first group of tests!)
```

<!-- grade_id: exercise_1_test_1, points: 1, grade: true, solution: false -->
```python
# Test cell: `exercise_1_test_1`
if False:
    for c in {'Y', 'K', 'M', 'S', 'B', 'D', 'V'}:
        from random import sample
        x = ''.join([sample('ACGT', 1)[0] for _ in range(2)])
        y = ''.join([sample('ACGT', 1)[0] for _ in range(2)])
        pattern = '{}{}{}'.format(x, c, y)
        ans = set(re.findall(bio_to_regex(pattern), dna_seq))
        print("assert set(re.findall(bio_to_regex('{}'), dna_seq)) == {}".format(pattern, ans))
        
assert set(re.findall(bio_to_regex('GABAT'), dna_seq)) == {'GACAT', 'GAGAT', 'GATAT'}
assert set(re.findall(bio_to_regex('GAVCA'), dna_seq)) == {'GACCA', 'GAACA'}
assert set(re.findall(bio_to_regex('TGYGG'), dna_seq)) == {'TGTGG', 'TGCGG'}
assert set(re.findall(bio_to_regex('GCKAA'), dna_seq)) == {'GCGAA'}
assert set(re.findall(bio_to_regex('ATSCA'), dna_seq)) == {'ATCCA'}
assert set(re.findall(bio_to_regex('GCMTT'), dna_seq)) == {'GCCTT', 'GCATT'}
assert set(re.findall(bio_to_regex('AGDCC'), dna_seq)) == {'AGTCC', 'AGACC'}

print("\n(Passed second set of tests!)")
```

Output:
```
(Passed second set of tests!)
```

**Restriction sites.** When an enzyme cuts the string, it does it in a certain location with respect to the target pattern. This information is encoded as a _restriction site_.

The way a biologist specifies the restriction site is with a special notation that embeds the cut in the pattern. For example, there is one enzyme that has a restriction site of the form, `ANT|AAT`, where the vertical bar, `'|'`, shows where the enzyme will split the sequence. So, if the input DNA sequence were

```
   GCATAGTAATGTATTAATGGC
```

then there would two matches:

```
   GCATAGTAATGTATTAATGGC
       ^^^^^^  ^^^^^^
       match!  match!
```

Furthermore, there would be two cuts, since this enzyme splits its pattern in the middle (between `ANT` and `AAT`):

```
   GCATAGT|AATGTATT|AATGGC
       ^^^ ^^^  ^^^ ^^^
```

That would result in three fragments: `GCATAGT`, `AATGTATT`, and `AATGGC`.

**Exercise 3** (5 points). Complete the function, `sim_cuts(site_pattern, s)`, below. The first argument, `site_pattern`, is the biologist's restriction site pattern, e.g., `ANT|AAT`, where there may be an embedded cut. The second argument, `s`, is the DNA sequence to cut. The function should return the fragments in the sequence order.

For the preceding example,

```python
  sim_cuts('ANT|AAT', 'GCATAGTAATGTATTAATGGC') == ['GCATAGT', 'AATGTATT', 'AATGGC']
```

> **Note.** There are *two* test cells, below. Both must pass for full credit, but if only one passes, you'll at least get some partial credit.

```python
s = "GCATAGTAATGTATTAATGGC"

site_pattern = "ANT|AAT"
```

```python
pattern = site_pattern.replace("|","")
pattern
```

Output:
```
'ANTAAT'
```

```python
pattern=bio_to_regex(pattern)
pattern
```

Output:
```
'A[ACGT]TAAT'
```

```python
match = re.search(pattern,s)
```

```python
match
```

Output:
```
<re.Match object; span=(4, 10), match='AGTAAT'>
```

```python
index_list = [0]
 
for match in re.finditer(pattern,s):
    index_list.append(match.start()+3)
    print(match)
    print(match.start())
```

Output:
```
<re.Match object; span=(4, 10), match='AGTAAT'>
4
<re.Match object; span=(12, 18), match='ATTAAT'>
12
```

```python
index_list
```

Output:
```
[0, 7, 15]
```

```python
len(s)
```

Output:
```
21
```

```python
index_list.append(len(s))
```

```python
index_list
```

Output:
```
[0, 7, 15, 21]
```

```python
len(index_list)
```

Output:
```
4
```

```python
fragments=[]

try:
    for i in range(len(index_list)):
        fragments.append(s[index_list[i]:index_list[i+1]])
except:
    pass
```

```python
fragments
```

Output:
```
['GCATAGT', 'AATGTATT', 'AATGGC']
```

```python
s[0:7]
```

Output:
```
'GCATAGT'
```

```python
s[7:15]
```

Output:
```
'AATGTATT'
```

```python
s[16:len(s)]
```

Output:
```
'ATGGC'
```

```python
s = "GCATAGTAATGTATTAATGGC"

site_pattern = "ANT|AAT"
```

```python
pattern = site_pattern.replace("|","")
pattern=bio_to_regex(pattern)
```

```python
pattern
```

Output:
```
'A[ACGT]TAAT'
```

```python
index_list = [0]
 
for match in re.finditer(pattern,s):
    index_list.append(match.start()+3)
```

```python
index_list
```

Output:
```
[0, 7, 15]
```

```python
fragments=[]

try:
    for i in range(len(index_list)):
        fragments.append(s[index_list[i]:index_list[i+1]])
except:
    pass

fragments.append(s[index_list[-1]:])
```

```python
fragments
```

Output:
```
['GCATAGT', 'AATGTATT', 'AATGGC']
```

<!-- grade_id: exercise_2, points: 0, grade: false, solution: true -->
```python
def sim_cuts(site_pattern, s):
    pattern = site_pattern.replace("|","")
    pattern=bio_to_regex(pattern)
    
    index_list = [0]
 
    for match in re.finditer(pattern,s):
        index_list.append(match.start()+3)
        
    fragments=[]

    try:
        for i in range(len(index_list)):
            fragments.append(s[index_list[i]:index_list[i+1]])
    except:
        pass

    fragments.append(s[index_list[-1]:])
    
    return fragments
```

<!-- grade_id: exercise_2_test_0, points: 3, grade: true, solution: false -->
```python
# Test cell: `exercise_3_test_0`

def check_sim_cuts(bio_pattern, s, true_cuts):
    print("\nChecking: '{}'...".format(bio_pattern))
    your_cuts = sim_cuts(bio_pattern, s)
    print("   Your result ({} fragments): {}".format(len(your_cuts), your_cuts))
    print("   True result ({}): {}".format(len(true_cuts), true_cuts))
    assert your_cuts == true_cuts, "Did not match!"
    print("   ==> Matched!")

# Check a simple case:
check_sim_cuts('ANT|AAT', 'GCATAGTAATGTATTAATGGC', ['GCATAGT', 'AATGTATT', 'AATGGC'])

print("\n(Passed first test of Exercise 3; two more to go in the next cell.)")
```

Output:
```
Checking: 'ANT|AAT'...
   Your result (3 fragments): ['GCATAGT', 'AATGTATT', 'AATGGC']
   True result (3): ['GCATAGT', 'AATGTATT', 'AATGGC']
   ==> Matched!

(Passed first test of Exercise 3; two more to go in the next cell.)
```

```python
len('ANT')
```

Output:
```
3
```

```python
s = 'ANGT|ANT'
parts = s.split('|')
parts
```

Output:
```
['ANGT', 'ANT']
```

```python
len(parts[0])
```

Output:
```
4
```

```python
def sim_cuts(site_pattern, s):
    
    s = 'ANGT|ANT'
    parts = s.split('|')
    parts
    d = len(parts[0])
    
    pattern = site_pattern.replace("|","")
    pattern=bio_to_regex(pattern)
    
    index_list = [0]
 
    for match in re.finditer(pattern,s):
        index_list.append(match.start()+d)
        
    fragments=[]

    try:
        for i in range(len(index_list)):
            fragments.append(s[index_list[i]:index_list[i+1]])
    except:
        pass

    fragments.append(s[index_list[-1]:])
    
    return fragments
```

## SIM_CUTS

```python
def sim_cuts(site_pattern, s):
    
    split_pattern = site_pattern.split('|')
    
    pattern_bio=''.join(split_pattern)
    
    m_iter=re.finditer(bio_to_regex(pattern_bio),s)
    
    offset = len(split_pattern[0]) # This line calculates the length of the first element of split_pattern
    
    inds = [0]+[m.span()[0] + offset for m in m_iter] + [len(s)] # m.span()[0] can be replaced by m.start()
    
    return(
    [''.join(s[start:end]) for start, end in zip(inds[:-1],inds[1:])]
    )
```

```python
inds = [0, 7, 14, 23]
```

```python
for i in zip(inds[:-1],inds[1:]):
    print(i)
```

Output:
```
(0, 7)
(7, 14)
(14, 23)
```

```python
for start, end in zip(inds[:-1],inds[1:]):
    print(start)
```

Output:
```
0
7
14
```

The `sim_cuts` function appears to be designed to find and split a given string `s` based on a simplified biological site pattern specified as `site_pattern`. The goal is to identify occurrences of the site pattern in the string `s`, and then split the string `s` at the positions where these site pattern occurrences start.

Here's a step-by-step explanation of the code:

1. `split_pattern = site_pattern.split('|')`: The `site_pattern` is a string that may contain one or more "|" (pipe) characters, which are used as delimiters. This line splits the `site_pattern` into a list called `split_pattern`, where each element of the list is a substring of `site_pattern` between "|" characters.

2. `pattern_bio = ''.join(split_pattern)`: This line joins the elements of the `split_pattern` list back together into a single string called `pattern_bio`. It effectively removes the "|" characters and creates a single string representing the combined pattern.

3. `m_iter = re.finditer(bio_to_regex(pattern_bio), s)`: This line uses the `bio_to_regex` function (which converts a biological pattern to a regular expression) to create a regular expression pattern from `pattern_bio`. It then uses `re.finditer` to find all non-overlapping matches of this pattern in the string `s`. The result is an iterator `m_iter` containing match objects.

4. `offset = len(split_pattern[0])`: This line calculates the length of the first element of `split_pattern` and assigns it to the variable `offset`. This offset is used to adjust the positions of the found matches to account for the length of the first element.

5. `inds = [0] + [m.span()[0] + offset for m in m_iter] + [len(s)]`: This line constructs a list `inds`, which contains the indices where the string `s` should be split. It starts with 0 (the beginning of the string), adds the starting positions of the matches found in `m_iter` (adjusted by the `offset`), and ends with the length of the string `s`.

6. `return (...)`: The function returns a list of substrings extracted from the original string `s` based on the calculated indices in `inds`. It uses list comprehension and the `zip` function to iterate through pairs of indices and extracts the corresponding substrings.

The result is a list of substrings from `s`, split at the positions where the specified site pattern occurs. This function is particularly useful in bioinformatics and sequence analysis to extract and analyze sequences around specific motifs or patterns.

```python
import re

# Define a regular expression pattern
pattern = r'\d+'  # This pattern matches one or more digits

# Input string
text = "There are 42 apples and 123 oranges on the table."

# Use re.finditer to find all matches of the pattern in the text
matches = re.finditer(pattern, text)

# Iterate through the matches and print their start and end positions
for match in matches:
    print("Match found:", match.group())  # match.group() returns the matched text
    print("Start position:", match.start())  # match.start() returns the start position
    print("End position:", match.end())      # match.end() returns the end position
    print("")

# Output:
# Match found: 42
# Start position: 11
# End position: 13
#
# Match found: 123
# Start position: 25
# End position: 28

```

Output:
```
Match found: 42
Start position: 10
End position: 12

Match found: 123
Start position: 24
End position: 27
```

```python
matches
```

Output:
```
<callable_iterator at 0x7f3270075490>
```

```python
my_string = "Hello, World!"
length = len(my_string)
print("Length of the string:", length)  # Output: 13

```

Output:
```
Length of the string: 13
```

```python
my_list = [1, 2, 3, 4, 5]
length = len(my_list)
print("Length of the list:", length)  # Output: 5

```

Output:
```
Length of the list: 5
```

```python
# Create two lists
list1 = [1, 2, 3]
list2 = ['a', 'b', 'c']

# Use zip() to combine the two lists
zipped = zip(list1, list2)

# Convert the result to a list (optional, for better visualization)
result = list(zipped)

# Print the zipped result
print(result)

```

Output:
```
[(1, 'a'), (2, 'b'), (3, 'c')]
```

```python
# Define a string
s = "Hello, World!"

# Specify start and end positions
start = 0  # Start from the beginning
end = 5    # End at position 5 (exclusive)

# Use ''.join(s[start:end]) to extract and concatenate characters
result = ''.join(s[start:end])

# Print the result
print(result)

```

Output:
```
Hello
```

<!-- grade_id: exercise_2_test_1, points: 2, grade: true, solution: false -->
```python
# Test cell: `exercise_2_test_1`

check_sim_cuts('ANT|AAT', dna_seq, ['ATGGCAATAACCCCCCGTTTCTACTTCTAGAGGAGAAAAGTATTGACATGAGCGCTCCCGGCACAAGGGCCAAAGAAGTCTCCAATTTCTTATTTCCGAATGACATGCGTCTCCTTGCGGGTAAATCACCGACCGCAATTCATAGAAGCCTGGGGGAACAGATAGGTCTAATTAGCTTAAGAGAGTAAATCCTGGGATCATTCAGTAGTAACCATAAACTTACGCTGGGGCTTCTTCGGCGGATTTTTACAGTTACCAACCAGGAGATTTGAAGTAAATCAGTTGAGGATTTAGCCGCGCTATCCGGTAATCTCCAAATTAAAACATACCGTTCCATGAAGGCTAGAATTACTTACCGGCCTTTTCCATGCCTGCGCTATACCCCCCCACTCTCCCGCTTATCCGTCCGAGCGGAGGCAGTGCGATCCTCCGTTAAGATATTCTTACGTGTGACGTAGCTATGTATTTTGCAGAGCTGGCGAACGCGTTGAACACTTCACAGATGGTAGGGATTCGGGTAAAGGGCGTATAATTGGGGACTAACATAGGCGTAGACTACGATGGCGCCAACTCAATCGCAGCTCGAGCGCCCTGAATAACGTACTCATCTCAACTCATTCTCGGCAATCTACCGAGCGACTCGATTATCAACGGCTGTCTAGCAGTTCTAATCTTTTGCCAGCATCGTAATAGCCTCCAAGAGATTGATGATAGCTATCGGCACAGAACTGAGACGGCGCCGATGGATAGCGGACTTTCGGTCAACCACAATTCCCCACGGGACAGGTCCTGCGGTGCGCATCACTCTGAATGTACAAGCAACCCAAGTGGGCCGAGCCTGGACTCAGCTGGTTCCTGCGTGAGCTCGAGACTCGGGATGACAGCTCTTTAAACATAGAGCGGGGGCGTCGAACGGTCGAGAAAGTCATAGTACCTCGGGTACCAACTTACTCAGGTTATTGCTTGAAGCTGTACTATTTTAGGGGGGGAGCGCTGAAGGTCTCTTCTTCTCATGACTGAACTCGCGAGGGTCGTGAAGTCGGTTCCTTCAATGGTTAAAAAACAAAGGCTTACTGTGCGCAGAGGAACGCCCATCTAGCGGCTGGCGTCTTGAATGCTCGGTCCCCTTTGTCATTCCGGATT',
 'AATCCATTTCCCTCATTCACGAGCTTGCGAAGTCTACATTGGTATATGAATGCGACCTAGAAGAGGGCGCTTAAAATTGGCAGTGGTTGATGCTCTAAACTCCATTTGGTTTACTCGTGCATCACCGCGATAGGCTGACAAAGGTTTAACATTGAATAGCAAGGCACTTCCGGTCTCAATGAACGGCCGGGAAAGGTACGCGCGCGGTATGGGAGGATCAAGGGGCCAATAGAGAGGCTCCTCTCTCACTCGCTAGGAGGCAAATGTAAAACAATGGTTACTGCATCGATACATAAAACATGTCCATCGGTTGCCCAAAGTGTTAAGTGTCTATCACCCCTAGGGCCGTTTCCCGCATATAAACGCCAGGTTGTATCCGCATTTGATGCTACCGTGGATGAGTCTGCGTCGAGCGCGCCGCACGAATGTTGCAATGTATTGCATGAGTAGGGTTGACTAAGAGCCGTTAGATGCGTCGCTGTACT',
 'AATAGTTGTCGACAGACCGTCGAGATTAGAAAATGGTACCAGCATTTTCGGAGGTTCTCTAACTAGTATGGATTGCGGTGTCTTCACTGTGCTGCGGCTACCCATCGCCTGAAATCCAGCTGGTGTCAAGCCATCCCCTCTCCGGGACGCCGCATGTAGTGAAACATATACGTTGCACGGGTTCACCGCGGTCCGTTCTGAGTCGACCAAGGACACAATCGAGCTCCGATCCGTACCCTCGACAAACTTGTACCCGACCCCCGGAGCTTGCCAGCTCCTCGGGTATCATGGAGCCTGTGGTTCATCGCGTCCGATATCAAACTTCGTCATGATAAAGTCCCCCCCTCGGGAGTACCAGAGAAGATGACTACTGAGTTGTGCGAT'])
check_sim_cuts('GCRW|TG', dna_seq, ['ATGGCAATAACCCCCCGTTTCTACTTCTAGAGGAGAAAAGTATTGACATGAGCGCTCCCGGCACAAGGGCCAAAGAAGTCTCCAATTTCTTATTTCCGAATGACATGCGTCTCCTTGCGGGTAAATCACCGACCGCAATTCATAGAAGCCTGGGGGAACAGATAGGTCTAATTAGCTTAAGAGAGTAAATCCTGGGATCATTCAGTAGTAACCATAAACTTACGCTGGGGCTTCTTCGGCGGATTTTTACAGTTACCAACCAGGAGATTTGAAGTAAATCAGTTGAGGATTTAGCCGCGCTATCCGGTAATCTCCAAATTAAAACATACCGTTCCATGAAGGCTAGAATTACTTACCGGCCTTTTCCATGCCTGCGCTATACCCCCCCACTCTCCCGCTTATCCGTCCGAGCGGAGGCAGTGCGATCCTCCGTTAAGATATTCTTACGTGTGACGTAGCTATGTATTTTGCAGAGCTGGCGAACGCGT',
 'TGAACACTTCACAGATGGTAGGGATTCGGGTAAAGGGCGTATAATTGGGGACTAACATAGGCGTAGACTACGATGGCGCCAACTCAATCGCAGCTCGAGCGCCCTGAATAACGTACTCATCTCAACTCATTCTCGGCAATCTACCGAGCGACTCGATTATCAACGGCTGTCTAGCAGTTCTAATCTTTTGCCAGCATCGTAATAGCCTCCAAGAGATTGATGATAGCTATCGGCACAGAACTGAGACGGCGCCGATGGATAGCGGACTTTCGGTCAACCACAATTCCCCACGGGACAGGTCCTGCGGTGCGCATCACTCTGAATGTACAAGCAACCCAAGTGGGCCGAGCCTGGACTCAGCTGGTTCCTGCGTGAGCTCGAGACTCGGGATGACAGCTCTTTAAACATAGAGCGGGGGCGTCGAACGGTCGAGAAAGTCATAGTACCTCGGGTACCAACTTACTCAGGTTATTGCTTGAAGCTGTACTATTTTAGGGGGGGAGCGCTGAAGGTCTCTTCTTCTCATGACTGAACTCGCGAGGGTCGTGAAGTCGGTTCCTTCAATGGTTAAAAAACAAAGGCTTACTGTGCGCAGAGGAACGCCCATCTAGCGGCTGGCGTCTTGAATGCTCGGTCCCCTTTGTCATTCCGGATTAATCCATTTCCCTCATTCACGAGCTTGCGAAGTCTACATTGGTATATGAATGCGACCTAGAAGAGGGCGCTTAAAATTGGCAGTGGTTGATGCTCTAAACTCCATTTGGTTTACTCGTGCATCACCGCGATAGGCTGACAAAGGTTTAACATTGAATAGCAAGGCACTTCCGGTCTCAATGAACGGCCGGGAAAGGTACGCGCGCGGTATGGGAGGATCAAGGGGCCAATAGAGAGGCTCCTCTCTCACTCGCTAGGAGGCAAATGTAAAACAATGGTTACTGCATCGATACATAAAACATGTCCATCGGTTGCCCAAAGTGTTAAGTGTCTATCACCCCTAGGGCCGTTTCCCGCATATAAACGCCAGGTTGTATCCGCATTTGATGCTACCGTGGATGAGTCTGCGTCGAGCGCGCCGCACGAATGTTGCAA',
 'TGTATTGCATGAGTAGGGTTGACTAAGAGCCGTTAGATGCGTCGCTGTACTAATAGTTGTCGACAGACCGTCGAGATTAGAAAATGGTACCAGCATTTTCGGAGGTTCTCTAACTAGTATGGATTGCGGTGTCTTCACTGTGCTGCGGCTACCCATCGCCTGAAATCCAGCTGGTGTCAAGCCATCCCCTCTCCGGGACGCCGCATGTAGTGAAACATATACGTTGCACGGGTTCACCGCGGTCCGTTCTGAGTCGACCAAGGACACAATCGAGCTCCGATCCGTACCCTCGACAAACTTGTACCCGACCCCCGGAGCTTGCCAGCTCCTCGGGTATCATGGAGCCTGTGGTTCATCGCGTCCGATATCAAACTTCGTCATGATAAAGTCCCCCCCTCGGGAGTACCAGAGAAGATGACTACTGAGTTGTGCGAT'])

print("\n(Passed second tests of Exercise 3!)")
```

Output:
```
Checking: 'ANT|AAT'...
   Your result (3 fragments): ['ATGGCAATAACCCCCCGTTTCTACTTCTAGAGGAGAAAAGTATTGACATGAGCGCTCCCGGCACAAGGGCCAAAGAAGTCTCCAATTTCTTATTTCCGAATGACATGCGTCTCCTTGCGGGTAAATCACCGACCGCAATTCATAGAAGCCTGGGGGAACAGATAGGTCTAATTAGCTTAAGAGAGTAAATCCTGGGATCATTCAGTAGTAACCATAAACTTACGCTGGGGCTTCTTCGGCGGATTTTTACAGTTACCAACCAGGAGATTTGAAGTAAATCAGTTGAGGATTTAGCCGCGCTATCCGGTAATCTCCAAATTAAAACATACCGTTCCATGAAGGCTAGAATTACTTACCGGCCTTTTCCATGCCTGCGCTATACCCCCCCACTCTCCCGCTTATCCGTCCGAGCGGAGGCAGTGCGATCCTCCGTTAAGATATTCTTACGTGTGACGTAGCTATGTATTTTGCAGAGCTGGCGAACGCGTTGAACACTTCACAGATGGTAGGGATTCGGGTAAAGGGCGTATAATTGGGGACTAACATAGGCGTAGACTACGATGGCGCCAACTCAATCGCAGCTCGAGCGCCCTGAATAACGTACTCATCTCAACTCATTCTCGGCAATCTACCGAGCGACTCGATTATCAACGGCTGTCTAGCAGTTCTAATCTTTTGCCAGCATCGTAATAGCCTCCAAGAGATTGATGATAGCTATCGGCACAGAACTGAGACGGCGCCGATGGATAGCGGACTTTCGGTCAACCACAATTCCCCACGGGACAGGTCCTGCGGTGCGCATCACTCTGAATGTACAAGCAACCCAAGTGGGCCGAGCCTGGACTCAGCTGGTTCCTGCGTGAGCTCGAGACTCGGGATGACAGCTCTTTAAACATAGAGCGGGGGCGTCGAACGGTCGAGAAAGTCATAGTACCTCGGGTACCA
```

**Fin!** If you've reached this point and all tests above pass, your biologist friend thanks you and you are ready to submit your solution to this problem. Don't forget to save you work prior to submitting.

Portions of this problem were inspired by a fun book called [Python for Biologists](https://pythonforbiologists.com/python-books).